# PH 306 – Computational Physics
## Assignment Driver Notebook

**Name:**

**Date:**

---

### Instructions
- Complete all required functions in **`assignment.py`**.
- Use this notebook to run checks and inspect intermediate values.
- GitHub Classroom grading runs against **`assignment.py`**, not notebook cell output.
- Run all cells before submitting to verify your implementation.

```python
# Numerical Imports
import numpy as np
```

---

## Problem 1

Implement `problem1_array_operations()`

It should return `(x, y)` where:

- `x` is 100 evenly spaced values from $0$ to $2\pi$ (inclusive)
- `y = sin(x)`

```python
def problem1_array_operations() -> tuple[np.ndarray, np.ndarray]:
    """Return x and y for Problem 1.

    - x: 100 evenly spaced values from 0 to 2*pi (inclusive)
    - y: sin(x)
    """
    raise NotImplementedError("Implement problem1_array_operations")
```

<!-- grade_id: problem1-solution, points: 0, grade: false, solution: true -->
```python
x, y = problem1_array_operations()
x[:5], y[:5]
```

<!-- grade_id: problem1-tests, points: 2, grade: true, solution: false -->
```python
### PUBLIC CHECKS (similar to autograder)
assert isinstance(x, np.ndarray), "x must be a NumPy array"
assert len(x) == 100, "x must contain 100 elements"
assert np.isclose(x[0], 0.0), "x must start at 0"
assert np.isclose(x[-1], 2 * np.pi), "x must end at 2π"
assert isinstance(y, np.ndarray), "y must be a NumPy array"
assert np.allclose(y[0], 0), "y[0] is not close to 0"
assert np.allclose(y[-1], 0), "y[-1] is not close to 0"
print("Problem 1 checks passed")
```

---

## Problem 2

Implement `problem2_numerical_integration()`

It should return `(result, error)` where `result` is the numerical value of
\begin{equation}
\int_{-\infty}^{+\infty} e^{-x^2} dx
\end{equation}
and `error` is SciPy's absolute error estimate.

*Hint:* The exact answer is $\sqrt{\pi}$.

```python
# Problem 2
def problem2_numerical_integration() -> tuple[float, float]:
    """Return (result, error) for integral of exp(-x^2) from -inf to +inf."""
    raise NotImplementedError("Implement problem2_numerical_integration")
```

<!-- grade_id: problem2-solution, points: 0, grade: false, solution: true -->
```python
result, error = problem2_numerical_integration()
result, error
```

<!-- grade_id: problem2-tests, points: 3, grade: true, solution: false -->
```python
### PUBLIC CHECKS (similar to autograder)
assert isinstance(result, float), "result must be a float"
assert isinstance(error, float), "error must be a float"
assert np.isclose(result, np.sqrt(np.pi), rtol=1e-6), \
    f"result={result} is not close to sqrt(π)={np.sqrt(np.pi)}"
print("Problem 2 checks passed")
```